# CHEM 1000 - Spring 2026
Prof. Geoffrey Hutchison, University of Pittsburgh

## Graded Homework 7

For this homework, we'll focus on:
- categories for differential equations
- first-order differential equations
---

As a reminder, you do not need to use Python to solve the problems. If you want, you can use other methods, just put your answers in the appropriate places.

To turn in, either download as Notebook (.ipynb) or Print to PDF and upload to Gradescope.

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name and collaborators (i.e., anyone you discussed this with) below:

```python
NAME = ""
COLLABORATORS = ""
```

### Categories of Differential Equations

Knowing the type of a differential equation gives insight into the techniques to solve it.

- Type (ordinary vs. partial) .. look at the type of derivative (e.g., does it use partial derivatives)
- Order .. look at the highest derivative (second, third, first, etc.)
- Linear vs. Nonlinear - the terms that involve 'y' or the function, do they only include derivatives and the function itself, or are there terms with '$y^2$' or other non-linear terms?
- Homogeneous vs. Non-homogeneous - a homogeneous diff. eq. can be separated with terms involving 'x' on one side, and terms involving the function (e.g., y, u, etc.) on the other side.

Categorize the following functions by assigning a string

One:
$$
\frac{dy}{dx} + 3y = 0
$$

Two:
$$
\frac{d^{2} y(x)}{d x^{2}}+2 \frac{d y(x)}{d x}+4 y(x)=0
$$

Three:
$$
L \frac{d^{2} u}{d x^{2}}+g \sin u=0
$$

Four:
$$
\frac{\partial^{2} u}{\partial x^{2}}+\frac{\partial^{2} u}{\partial y^{2}}=0
$$

Five (the classical heat equation):
$$
\frac{\partial u(x, t)}{\partial t}=\kappa \frac{\partial^{2} u(x, t)}{\partial x^{2}}
$$

<!-- grade_id: categories-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
categories_one = 'linear homogeneous first order ordinary differential equation'
categories_two = '.. differential equation'
categories_three = '.. differential equation'
categories_four = '.. differential equation'
categories_five = '.. differential equation'
```

<!-- grade_id: categories-validation, points: 7, grade: true, solution: false -->
```python
# these should NOT be in the answer
assert('nonlinear' not in categories_one)
assert('non-homogeneous' not in categories_one)
assert('second' not in categories_one)
assert('partial' not in categories_one)
# these should be in the answer
assert('linear' in categories_one)
assert('homogeneous' in categories_one)
assert('first' in categories_one)
assert('ordinary' in categories_one)
assert('differential' in categories_one)

# more (hidden) tests for two through five

```

### Solving Differential Equations

We'll go through two examples a homogeneous first-order equation and then a non-homogeneous equation:

Solve:
$$
\frac{dy}{dx} + 3y = 0
$$

<!-- grade_id: cell-homogeneous, points: 5, grade: true, solution: true -->
```python
# YOUR ANSWER HERE

```

Show that your solution works by plugging it back into the differential equation and showing that it equals zero.

<!-- grade_id: cell-validate, points: 1, grade: true, solution: true -->
```python
# YOUR ANSWER HERE

```

Solve the following first-order linear non-homogenous ordinary differential equation:

$$
\frac{d y(x)}{d x}+\frac{2}{x} y(x)=x^{2}+2
$$

<!-- grade_id: cell-non-homogeneous, points: 5, grade: true, solution: true -->
```python
# YOUR ANSWER HERE

```

Show that your solution works by plugging it back into the differential equation

<!-- grade_id: validation-two, points: 2, grade: true, solution: true -->
```python
# YOUR ANSWER HERE

```